# Pandas NBgrader exercises

In this notebook you will find the 6 deliverable exercises regarding pandas. This notebook is autograded so please dont modify any cell except to add your solution. Modifying this notebook may result in you not getting a grade.

```python
import pandas as pd
import numpy as np
```

```python
data = {'animal': ['cat', 'cat', 'snake', 'dog', 'dog', 'cat', 'snake', 'cat', 'dog', 'dog'],
        'age': [2.5, 3, 0.5, np.nan, 5, 2, 4.5, np.nan, 7, 3],
        'visits': [1, 3, 2, 3, 2, 3, 1, 1, 2, 1],
        'priority': ['yes', 'yes', 'no', 'yes', 'no', 'no', 'no', 'yes', 'no', 'no']}

labels = ['a', 'b', 'c', 'd', 'e', 'f', 'g', 'h', 'i', 'j']


df = pd.DataFrame(data, index=labels)
```

**13.** Change the age in row 'f' to 1.5.

<!-- grade_id: cell-f8d4a64dd0fb6112, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
df.loc['f','age'] = 1.5
```

<!-- grade_id: cell-675b0274a23f22c9, points: 1, grade: true, solution: false -->
```python
assert df['age']['f'] == 1.5
```

**18.** Sort `df` first by the values in the 'age' in *decending* order, then by the value in the 'visit' column in *ascending* order.

<!-- grade_id: cell-bb44e438722da485, points: 0, grade: false, solution: true -->
```python
df = df.sort_values(by='visits')                    # FIRST we sort by 'visits'
df = df.sort_values(by='age', ascending = False)    # THEN  we sort by 'age'
                                                    # that's how we get it sorted by 'age' first and 'age' second
# later i saw it's possible this way:
# df.sort_values(['age', 'visits'], ascending=[False, True])

```

<!-- grade_id: cell-9fa22615a21cf3f1, points: 1, grade: true, solution: false -->
```python
b = True

for i in range(df.shape[0]-1):
    if df['age'].values[i] < df['age'].values[i+1]:
        b = False
        break
    
assert b
   
```

**23.** Given a DataFrame of numeric values, say
```python
df = pd.DataFrame(np.random.random(size=(5, 3))) # a 5x3 frame of float values
```

how do you subtract the row mean from each element in the row?

<!-- grade_id: cell-5137285459670e80, points: 0, grade: false, solution: true -->
```python
df = pd.DataFrame(np.random.random(size=(5, 3))) # a 5x3 frame of float values
result = df - df.describe().loc['mean']
print(result)

# later on I've found this solution on peers review, I find it smarter ;)
# result = df.sub(df.mean(axis=1), axis=0)
```

Output:
```
0         1         2
0  0.017135  0.158370 -0.181119
1  0.077734 -0.406285 -0.218716
2 -0.238105  0.194084 -0.004555
3 -0.233494 -0.023479  0.566463
4  0.376729  0.077309 -0.162073
```

<!-- grade_id: cell-8ceb24481d7d7a28, points: 2, grade: true, solution: false -->
```python
assert df.shape == (5,3)
assert df[0].max() <= 1
assert df[0].min() >= -1
```

**24.** Suppose you have DataFrame with 10 columns of real numbers, for example:

```python
df = pd.DataFrame(np.random.random(size=(5, 10)), columns=list('abcdefghij'))
```
Which column of numbers has the smallest sum? (Find that column's label.)

<!-- grade_id: cell-7830a1e1958310b9, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
df = pd.DataFrame(np.random.random(size=(5, 10)), columns=list('abcdefghij'))
idx = df.sum().idxmin()
```

<!-- grade_id: cell-a0daad4562804d2a, points: 2, grade: true, solution: false -->
```python
assert idx in 'abcdefghij'
```

The next three puzzles are slightly harder...

**26.** You have a DataFrame that consists of 10 columns of floating--point numbers. Suppose that exactly 5 entries in each row are NaN values. For each row of the DataFrame, find the *column* which contains the *third* NaN value.

(You should return a Series of column labels.)

<!-- grade_id: cell-e4f640a69590b406, points: 0, grade: false, solution: true -->
```python
data = {
    1:[np.nan,2,np.nan,4,5],
    2:[np.nan,2,3,np.nan,5],
    3:[np.nan,2,np.nan,4,5],
    4:[1,np.nan,3,np.nan,np.nan],
    5:[1,2,np.nan,4,5],
    6:[1,np.nan,3,np.nan,np.nan],
    7:[1,2,3,4,5],
    8:[1,np.nan,3,np.nan,np.nan],
    9:[1,2,3,4,np.nan],
    10:[1,np.nan,3,4,5]
}
# YOUR CODE HERE
df = pd.DataFrame(data)
tf_frame = df.isna()
cumsum_frame = tf_frame.cumsum(axis=1)
res = (cumsum_frame == 3).idxmax(axis=1)
indexes = list(res)
```

<!-- grade_id: cell-182477a2beb50458, points: 3, grade: true, solution: false -->
```python
assert indexes == [3, 8, 5, 6, 8]
assert res.shape == (5,)
assert list(res.values) == [3, 8, 5, 6, 8]
```

**27.** A DataFrame has a column of groups 'grps' and and column of numbers 'vals'. For example: 

```python
df = pd.DataFrame({'grps': list('aaabbcaabcccbbc'), 
                   'vals': [12,345,3,1,45,14,4,52,54,23,235,21,57,3,87]})
```
For each *group*, find the sum of the three greatest values.

<!-- grade_id: cell-da30c3e27df05854, points: 0, grade: false, solution: true -->
```python
df = pd.DataFrame({'grps': list('aaabbcaabcccbbc'), 
                   'vals': [12,345,3,1,45,14,4,52,54,23,235,21,57,3,87]})
# YOUR CODE HERE
res = df.groupby("grps").vals.nlargest(3).sum(level=0)
```

<!-- grade_id: cell-dd5c34897ebaa573, points: 3, grade: true, solution: false -->
```python
assert res.values[0] == 409
assert res.values[1] == 156
assert res.values[2] == 345
```